---
rule_files:
  # See alerts_validation_spec.rb for details of how stdin gets set:
  - fixtures/rules.yml

evaluation_interval: 1m

tests:
  - interval: 5m
    input_series:
      # Cells whose available memory craters
      # after the first two hours
      - series: 'firehose_value_metric_rep_capacity_remaining_memory{environment="test", bosh_job_id="1", bosh_job_name="diego-cells"}'
        values: '80000-10x24 30000-10x36'

      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_id="1", bosh_job_name="diego-cells"}'
        values: '100000-0x60'

      - series: 'firehose_value_metric_rep_capacity_remaining_memory{environment="test", bosh_job_id="2", bosh_job_name="diego-cells"}'
        values: '70000-10x60 20000-10x36'

      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_id="2", bosh_job_name="diego-cells"}'
        values: '100000-0x60'

      # Isolation segment whose memory stays stable throughout
      - series: 'firehose_value_metric_rep_capacity_remaining_memory{environment="test", bosh_job_id="1", bosh_job_name="diego-cells-iso-seg"}'
        values: '90000-0x60'

      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_id="1", bosh_job_name="diego-cells-iso-seg"}'
        values: '100000-0x60'

    alert_rule_test:
      # Does not fire when the percentage of free memory
      # is above the threshold
      - eval_time: 2h
        alertname: DiegoCellRepsReachingTotalMemoryCapacity

      # Does not fire when the percentage of free memory
      # has not been below the threshold for long enough
      - eval_time: 4h
        alertname: DiegoCellRepsReachingTotalMemoryCapacity

      # Fires when the percentage of free memory
      # has been below the threshold for 2 hours
      - eval_time: 5h
        alertname: DiegoCellRepsReachingTotalMemoryCapacity
        exp_alerts:
          - exp_labels:
              severity: warning
              environment: test
            exp_annotations:
              summary: Reps low free memory capacity
              bosh_job_name: diego-cells
              description: >
                Low free memory 75% for the total advertised rep memory capacity in the last 2 hours on average - if we lost an AZ we may not be able to find space for all the app instances we're currently running. Review if we need to scale.
              url: https://team-manual.cloud.service.gov.uk/architecture_decision_records/ADR021-cell-capacity-assignment-2/#decision

    promql_expr_test:
      - expr: rep_memory_capacity_pct:avg5m
        eval_time: 0m
        exp_samples:
          - labels: 'rep_memory_capacity_pct:avg5m{environment="test", bosh_job_name="diego-cells"}'
            value: 75
          - labels: 'rep_memory_capacity_pct:avg5m{environment="test", bosh_job_name="diego-cells-iso-seg"}'
            value: 90

  - interval: 5m
    input_series:
      # Each series has a fixed amount of memory for 1 hour
      # Two diego-cells series, and one diego-cells-iso-seg series.
      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_name="diego-cells", bosh_job_index="1"}'
        values: 100+0x12
      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_name="diego-cells", bosh_job_index="2"}'
        values: 100+0x12
      - series: 'firehose_value_metric_rep_capacity_total_memory{environment="test", bosh_job_name="diego-cells-iso-seg", bosh_job_index="1"}'
        values: 100+0x12
    promql_expr_test:
      - expr: diego_cells_deployed
        eval_time: 5m
        exp_samples:
          - labels: 'diego_cells_deployed{bosh_job_name="diego-cells", environment="test"}'
            value: 2
          - labels: 'diego_cells_deployed{bosh_job_name="diego-cells-iso-seg", environment="test"}'
            value: 1
